# Robotic Systems I (ECE-DK808)

## Electrical and Computer Engineering Department, University of Patras, Greece

**Instructor:** Konstantinos Chatzilygeroudis (costashatz@upatras.gr)

## Lab 7

### Gradients, Jacobians and Hessians

Let's start by computing the gradient and the derivative (Leibniz's notation) of a scalar function with vector-valued inputs. We have this nice quadratic:

$f(\boldsymbol{x}) = \frac{1}{2}\boldsymbol{x}^T\boldsymbol{Q}\boldsymbol{x}$

where $\boldsymbol{x}\in{\mathbb{R}^{3\times 1}}$ (column vector).

Let's begin by writing a Python function, `f()`, that computes the value:

```python
# Imports
import numpy as np # Linear Algebra
import scipy # For matrix exponential
import matplotlib.pyplot as plt # Plotting
```

<!-- grade_id: cell-fbbdeef44460b9ac, points: 0, grade: false, solution: true -->
```python
Q = np.diag([0.1, 0.5, 1.]) # Global var for ease of use

def f(x):
    ### TO-DO: Write the function! It should return a scalar!
    ### ANSWER: Insert code here
    return 0.5 * x.T @ Q @ x
    ### END of ANSWER
```

<!-- grade_id: cell-3f5cff008c4df0c2, points: 3, grade: true, solution: false -->
```python
x = np.array([[2., 1., 2.]]).T
assert(np.isclose(f(x), 2.45))

x = np.array([[-2., -1., -2.]]).T
assert(np.isclose(f(x), 2.45))

x = np.array([[-2., -1., 2.]]).T
assert(np.isclose(f(x), 2.45))

x = np.array([[-2., -1., 3.]]).T
assert(np.isclose(f(x), 4.95))

x = np.array([[-5., -1., 3.]]).T
assert(np.isclose(f(x), 6.))
```

The gradient of $f$ is denoted as $\nabla_{\boldsymbol{x}}f\in\mathbb{R}^{3\times 1}$ (column vector). Let's write a function in Python, `gradf()`, that computes this:

<!-- grade_id: cell-1d0102c082e5155d, points: 0, grade: false, solution: true -->
```python
def gradf(x):
    ### TO-DO: Write the gradient of f
    ### ANSWER: Insert code here
    return 0.5 * (Q + Q.T) @ x
    ### END of ANSWER
```

On the other hand, the derivative of $f$ is denoted as $\frac{\partial f}{\partial\boldsymbol{x}}\in\mathbb{R}^{1\times 3}$ (row vector). Let's write a function in Python, `df()`, that computes this:

<!-- grade_id: cell-95a35fba388dab76, points: 0, grade: false, solution: true -->
```python
def df(x):
    ### TO-DO: Write the derivative of f
    ### ANSWER: Insert code here
    return 0.5 * x.T @ (Q + Q.T)
    ### END of ANSWER
```

Are the two functions computing the same thing!? Let's find out:

```python
x0 = np.array([[-2., 1., 4.]]).T

gf = gradf(x0)
dfdx = df(x0)

print(gf)
print(dfdx)
```

Output:
```
[[-0.2]
 [ 0.5]
 [ 4. ]]
[[-0.2  0.5  4. ]]
```

What is the difference!? Basically, $\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_0) = \frac{\partial f}{\partial\boldsymbol{x}}\Big|_{\boldsymbol{x}=\boldsymbol{x}_0}^T$. Let's verify this!

<!-- grade_id: cell-339d2dc7b5e7e871, points: 3, grade: true, solution: false -->
```python
print(np.linalg.norm(gf-dfdx.T))

assert(np.isclose(gf, dfdx.T).all())

for _ in range(20):
    x0 = np.random.randn(3, 1)
    gf = gradf(x0)
    dfdx = df(x0)
    assert(np.isclose(gf, dfdx.T).all())

x = np.array([[2., 1., 2.]]).T
assert(np.isclose(df(x), np.array([[0.2, 0.5, 2.]])).all())

x = np.array([[-2., -1., -2.]]).T
assert(np.isclose(df(x), np.array([[-0.2, -0.5, -2.]])).all())

x = np.array([[-2., -1., 2.]]).T
assert(np.isclose(df(x), np.array([[-0.2, -0.5, 2.]])).all())

x = np.array([[-2., -1., 3.]]).T
assert(np.isclose(df(x), np.array([[-0.2, -0.5, 3.]])).all())

x = np.array([[-5., -1., 3.]]).T
assert(np.isclose(df(x), np.array([[-0.5, -0.5, 3.]])).all())
```

Output:
```
0.0
```

Now let's compute the Hessian of `f()` compactly:

<!-- grade_id: cell-2f047113f5e55385, points: 0, grade: false, solution: true -->
```python
def hessf(x):
    ### TO-DO: Write the Hessian of f
    ### ANSWER: Insert code here
    return 0.5 * (Q + Q.T)
    ### END of ANSWER
```

<!-- grade_id: cell-72cb96675671a965, points: 3, grade: true, solution: false -->
```python
x0 = np.random.randn(3, 1)
x1 = np.random.randn(3, 1)
assert(np.isclose(hessf(x0), hessf(x1)).all())
```

Now let's create a vector-valued function $f(\cdot):\mathbb{R}^N\to\mathbb{R}^M$. Let's do the **uncontrolled** pendulum:

```python
def pendulum(x, l = 1., g = 9.81):
    theta_ddot = (-g/l)*np.sin(x[:1])

    return np.concatenate([x[1:], theta_ddot], axis=0)
```

Let's write the Jacobian of this. But first, let's remind ourselves what the Jacobian is:

For $\boldsymbol{y} = f(\boldsymbol{x}) : \mathbb{R}^N\to\mathbb{R}^M$, we have:

$\frac{\partial f}{\partial\boldsymbol{x}} = \begin{bmatrix}\frac{\partial\boldsymbol{y}_1}{\partial\boldsymbol{x}_1} & \frac{\partial\boldsymbol{y}_1}{\partial\boldsymbol{x}_2} & \dots & \frac{\partial\boldsymbol{y}_1}{\partial\boldsymbol{x}_N}\\\frac{\partial\boldsymbol{y}_2}{\partial\boldsymbol{x}_1} & \frac{\partial\boldsymbol{y}_2}{\partial\boldsymbol{x}_2} & \dots & \frac{\partial\boldsymbol{y}_2}{\partial\boldsymbol{x}_N}\\\ddots & \ddots & \ddots & \ddots\\\frac{\partial\boldsymbol{y}_M}{\partial\boldsymbol{x}_1} & \frac{\partial\boldsymbol{y}_M}{\partial\boldsymbol{x}_2} & \dots & \frac{\partial\boldsymbol{y}_M}{\partial\boldsymbol{x}_N}\end{bmatrix}\in\mathbb{R}^{M\times N}$

<!-- grade_id: cell-875c89244047f644, points: 0, grade: false, solution: true -->
```python
def pendulum(x, l = 1., g = 9.81):
    theta_ddot = (-g/l)*np.sin(x[:1])

    return np.concatenate([x[1:], theta_ddot], axis=0)

def deriv_pendulum(x, l = 1., g = 9.81):
    ### TO-DO: Write the Jacobian of `pendulum()`
    ### ANSWER: Insert code here
    dx = np.block([[0., 1.], [(-g/l)*np.cos(x[:1]), 0.]])
    ### END of ANSWER
    return dx
```

<!-- grade_id: cell-8561c4fdd4f6c684, points: 3, grade: true, solution: false -->
```python
# Let's test it
x0 = np.array([[0.1, 0.]]).T
dx0 = np.array([[0., 1.], [-9.76099086, 0.]])
assert(np.isclose(deriv_pendulum(x0), dx0).all())

x0 = np.array([[0.1, 10.]]).T
assert(np.isclose(deriv_pendulum(x0), dx0).all())

x0 = np.array([[-20., 0.]]).T
dx0 = np.array([[0., 1.], [-4.00328503, 0.]])
print(deriv_pendulum(x0))
assert(np.isclose(deriv_pendulum(x0), dx0).all())
```

Output:
```
[[ 0.          1.        ]
 [-4.00328503  0.        ]]
```

### Newton's Method

Let's remember the Newton's method for **root finding**:

1) We first linearize around the current estimate $\boldsymbol{x}_k$:

$f(\boldsymbol{x}_k + \Delta\boldsymbol{x})\approx f(\boldsymbol{x}_k) + \frac{\partial f}{\partial\boldsymbol{x}}\Big|_{\boldsymbol{x}_k}\Delta\boldsymbol{x}$

2) We then set $f(\boldsymbol{x}_k + \Delta\boldsymbol{x}) = 0$ and solve for $\Delta\boldsymbol{x}$:

$f(\boldsymbol{x}_k) + \frac{\partial f}{\partial\boldsymbol{x}}\Big|_{\boldsymbol{x}_k}\Delta\boldsymbol{x} = 0$

$\Delta\boldsymbol{x} = -\Big(\frac{\partial f}{\partial\boldsymbol{x}}\Big|_{\boldsymbol{x}_k}\Big)^{-1}f(\boldsymbol{x}_k)$

3) We apply the correction $\boldsymbol{x}_{k+1} = \boldsymbol{x}_k + \Delta\boldsymbol{x}$

4) Repeat until convergence

Let's try the newton method on our `pendulum()`!

<!-- grade_id: cell-f212aa03741da02e, points: 0, grade: false, solution: true -->
```python
# Initial guess
x0 = np.array([[0.5, 0.]]).T

# Step for newton (we start at a big one to enable the while loop)
dx = np.ones((2,1))

def newton_root_finding_step(x):
    ### TO-DO: Write Newton's method step for Root Finding.
    ### You should compute the Δx (store it in a variable named `dx`)
    ### ANSWER: Insert code here
    dx = - np.linalg.inv(deriv_pendulum(x)) @ pendulum(x)
    ### END of ANSWER
    return x + dx, dx

# iterate
x = np.copy(x0)
iters = 0
while np.linalg.norm(dx) > 1e-6:
    x, dx = newton_root_finding_step(x)
    iters += 1

print("Found x:", x.T, "->", pendulum(x).T, "in", iters, "iterations!")
```

<!-- grade_id: cell-ab1c8a43b4dcdeab, points: 3, grade: true, solution: false -->
```python
x0 = np.array([[0.5, 0.]]).T
x, dx = newton_root_finding_step(x0)
assert(np.isclose(dx, np.array([[-0.54630249], [0.]]), rtol=1e-3).all())

x0 = np.array([[np.pi - 0.1, 0.]]).T
x, dx = newton_root_finding_step(x0)
assert(np.isclose(dx, np.array([[0.1], [0.]]), rtol=1e-3, atol=1e-3).all())
```

### Newton's Method for Optimization

Let's optimize `f()`! How can we do that with Newton's Method? We frame the optimization problem as a root finding problem at the gradient of $f$! Let's remember the procedure:

$\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k + \Delta\boldsymbol{x}) \approx \nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k) + \frac{\partial}{\partial\boldsymbol{x}}\Big(\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k)\Big)\Delta\boldsymbol{x} = 0$

$\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k + \Delta\boldsymbol{x}) \approx \nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k) + \nabla^2_{\boldsymbol{x}}f(\boldsymbol{x}_k)\Delta\boldsymbol{x} = 0$

$\Delta\boldsymbol{x} = -\Big(\nabla^2_{\boldsymbol{x}}f(\boldsymbol{x}_k)\Big)^{-1}\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k)$

$\boldsymbol{x}_{k+1} = \boldsymbol{x}_k + \Delta\boldsymbol{x}$

Cool! Let's do this on `f()`:

<!-- grade_id: cell-47571e84440a86a6, points: 0, grade: false, solution: true -->
```python
# Initial point
x0 = np.array([[-2., 1., 4.]]).T

# Step for newton (we start at a big one to enable the while loop)
dx = np.ones((3,1))

def newton_step(x):
    ### TO-DO: Write Newton's method step for Optimization
    ### You should compute the Δx (store it in a variable named `dx`)
    ### Do not forget to update x with Δx!
    ### ANSWER: Insert code here
    dx = - np.linalg.inv(hessf(x)) @ gradf(x)
    ### END of ANSWER
    return x + dx, dx

# iterate
x = np.copy(x0)
iters = 0
while np.linalg.norm(dx) > 1e-6:
    x, dx = newton_step(x)
    iters += 1

print("Found x:", x.T, "->", f(x), "in", iters, "iterations!")
```

<!-- grade_id: cell-162206bf77bb67c6, points: 3, grade: true, solution: false -->
```python
x0 = np.array([[-2., 1., 4.]]).T
x, dx = newton_step(x0)
assert(np.isclose(dx, -x0, rtol=1e-3).all())

x0 = np.array([[20., 10., 40.]]).T
x, dx = newton_step(x0)
assert(np.isclose(dx, -x0, rtol=1e-3, atol=1e-3).all())

for _ in range(100):
    x0 = np.random.randn(3, 1)
    x, dx = newton_step(x0)
    assert(np.isclose(dx, -x0, rtol=1e-3, atol=1e-3).all())
```

Wow! This was fast convergence! Only **2** iterations! Try putting the initial point far away! In convex functions (like our function), Newton's method has **quadratic convergence rate**: this basically means **blazingly fast!!** Notice also that we used the autodiff gradients/Hessian from *jax*!

### Practical Newton's Method (with tricks)

In practice, Newton's method can *maximize* instead of minimizing and also overshoot the local minimum, oscillate or even explode! Let's start optimizing a harder function and see how we can address those issues of Newton's method.

Let's try to optimize the [Rastrigin Function](https://www.sfu.ca/~ssurjano/rastr.html) in 2D. Let's plot it to see how it looks:

```python
def f(x):
    A = 10.
    d = x.shape[0]
    return A * d + np.sum(x.T @ x) - A * np.sum(np.cos(2. * np.pi * x))

plt.close()
fig = plt.figure()
ax = fig.add_subplot(1, 1, 1)

N = 80
x1 = np.linspace(-2., 2., N)
x2 = np.linspace(-2., 2., N)

X, Y = np.meshgrid(x1, x2)

val = np.zeros((N, N))
for i in range(N):
    for j in range(N):
        xx = np.zeros((2, 1))
        xx[0] = X[i, j]
        xx[1] = Y[i, j]
        val[i, j] = f(xx)

CS = ax.contour(x1.reshape((N,)), x2.reshape((N,)), val)
```

Nice function with many many local minima! Let's use Newton's method on this!

<!-- grade_id: cell-7563568f6cf609d8, points: 0, grade: false, solution: true -->
```python
# First let's compute its derivatives!
def df_dx(x):
    A = 10.
    ### TO-DO: Compute and return the derivative of f(x) (Rastrigin)
    ### ANSWER: Insert code here
    return 2 * x.T + A * 2 * np.pi * np.sin(2*np.pi*x.T)
    ### END of ANSWER

def ddf_ddx(x):
    A = 10.
    ### TO-DO: Compute and return the Hessia n of f(x) (Rastrigin)
    ### ANSWER: Insert code here
    a = np.block([[A * 4 * np.pi**2 * np.cos(2*np.pi*x[0]), 0.], [0., A * 4 * np.pi**2 * np.cos(2*np.pi*x[1])]])
    H = 2 * np.eye(2) + a
    return H
    ### END of ANSWER
```

<!-- grade_id: cell-7a998a29bbb722d4, points: 3, grade: true, solution: false -->
```python
def finite_diff(f, z, M, eps = 1e-4):
    N = z.shape[0]
    jac = np.zeros((M, N))
    v = np.zeros((N, 1))
    for i in range(N):
        zp = np.copy(z)
        v[i] = eps
        zp = zp + v
        zm = np.copy(z)
        v[i] = -eps
        zm = zm + v
        jac[:, i:i+1] = (((f(zp) - f(zm))) / (2. * eps)).reshape((-1, 1))
    return jac

for _ in range(100):
    x = np.random.randn(2, 1)
    assert(np.isclose(finite_diff(f, x, 1), df_dx(x), rtol=1e-3, atol=1e-3).all())
    assert(np.isclose(finite_diff(df_dx, x, x.shape[0]), ddf_ddx(x), rtol=1e-3, atol=1e-3).all())
```

```python
# Initial point
x0 = np.array([[1., 0.7]]).T
# x0 = np.array([[1., 1.25]]).T # Uncomment and watch it explode!

ax.plot(x0[0, 0], x0[1, 0], 'rx')

print("Initial value:", f(x0))

fig # show figure again with updated point(s)
```

<!-- grade_id: cell-623ca4274e21b479, points: 0, grade: false, solution: true -->
```python
def newton_step(x):
    ### TO-DO: Write Newton's method step for Optimization
    ### Compute Δx (store it in a variabe named `dx`)
    ### ANSWER: Insert code here
    dx = - np.linalg.inv(ddf_ddx(x)) @ df_dx(x).T
    ### END of ANSWER
    return x + dx
```

<!-- grade_id: cell-adbec1b5ecdcda7a, points: 3, grade: true, solution: false -->
```python
np.random.seed(10)

x = np.random.randn(2, 1)
xn = newton_step(x)
assert(np.isclose(xn, np.array([[1.63122303, -0.0027116]]).T, rtol=1e-3, atol=1e-3).all())

x = np.random.randn(2, 1)
xn = newton_step(x)
assert(np.isclose(xn, np.array([[-1.50668149, 0.]]).T, rtol=1e-3, atol=1e-3).all())

x = np.random.randn(2, 1)
xn = newton_step(x)
assert(np.isclose(xn, np.array([[0.47265641, 0.11992915]]).T, rtol=1e-3, atol=1e-3).all())

np.random.seed(None)
```

```python
# Let's start solving
x = np.copy(x0)
```

```python
# iterate
x = newton_step(x)
print("x:", x.T, "->", f(x))

ax.plot(x[0, 0], x[1, 0], 'rx')

fig # show figure again with updated point(s)
```

We find a local maximum! **AND** we can easily explode! Let's fix that with **damping** and **Line search**. Let's remind ourselves of the damping procedure (this helps to always minimize and regularize the optimizer):

$\boldsymbol{H} = \nabla^2_{\boldsymbol{x}}f(\boldsymbol{x}_k)$

$\text{while } \boldsymbol{H}\preceq 0:$

$\quad\boldsymbol{H} = \boldsymbol{H} + \beta\boldsymbol{I}$

$\Delta\boldsymbol{x} = -\boldsymbol{H}^{-1}\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k)$

$\boldsymbol{x}_{k+1} = \boldsymbol{x}_k + \Delta\boldsymbol{x}$

Let's create a function that does the **damped Netwon's iterate** (aka one step):

<!-- grade_id: cell-fa14d7c52e8994b8, points: 0, grade: false, solution: true -->
```python
def is_pos_def(x):
    return np.all(np.linalg.eigvals(x) > 0)

def damped_newton_step(x, beta = 100.):
    ### TO-DO: Implement Damped Newton step
    ### Compute Δx (store it in a variabe named `dx`)
    ### ANSWER: Insert code here
    H = ddf_ddx(x)
    while(not is_pos_def(H)):
        H = H + beta * np.eye(2)
    dx = - np.linalg.inv(H) @ df_dx(x).T
    ### END of ANSWER
    return x + dx
```

<!-- grade_id: cell-84e29371f4464e7b, points: 3, grade: true, solution: false -->
```python
np.random.seed(10)

x = np.random.randn(2, 1)
xn = damped_newton_step(x)
assert(np.isclose(xn, np.array([[-5.52664708, 1.22929842]]).T, rtol=1e-3, atol=1e-3).all())

x = np.random.randn(2, 1)
xn = damped_newton_step(x)
assert(np.isclose(xn, np.array([[-2.17512681, -0.00420789]]).T, rtol=1e-3, atol=1e-3).all())

x = np.random.randn(2, 1)
xn = damped_newton_step(x)
assert(np.isclose(xn, np.array([[3.177012, -0.98422351]]).T, rtol=1e-3, atol=1e-3).all())

np.random.seed(None)
```

Let's optimize the Rastrigin function with our new function:

```python
# Plot the function!
plt.close()
fig = plt.figure()
ax = fig.add_subplot(1, 1, 1)

N = 80
x1 = np.linspace(-3., 3., N)
x2 = np.linspace(-3., 3., N)

X, Y = np.meshgrid(x1, x2)

val = np.zeros((N, N))
for i in range(N):
    for j in range(N):
        xx = np.zeros((2, 1))
        xx[0] = X[i, j]
        xx[1] = Y[i, j]
        val[i, j] = f(xx)

CS = ax.contour(x1.reshape((N,)), x2.reshape((N,)), val)

# Initial point
x0 = np.array([[1., 0.7]]).T
# x0 = np.array([[1., 1.25]]).T # Uncomment and watch it explode!

ax.plot(x0[0, 0], x0[1, 0], 'rx')

print("Initial value:", f(x0))

x = np.copy(x0)
```

```python
# iterate
x = damped_newton_step(x)

print("x:", x.T, "->", f(x))

ax.plot(x[0, 0], x[1, 0], 'rx')

fig # show figure again with updated point(s)
```

We are still overshooting and we can explode! We need **line search**! Let's remember the Armijo rule:

$\alpha = 1$

$\text{while } f(\boldsymbol{x}_k + \alpha\Delta\boldsymbol{x}) > f(\boldsymbol{x}_k) + b\alpha\nabla_{\boldsymbol{x}}f(\boldsymbol{x}_k)^T\Delta\boldsymbol{x}:$

$\quad\alpha = c\alpha$

$\boldsymbol{x}_{k+1} = \boldsymbol{x}_k + \alpha\Delta\boldsymbol{x}$

where $0<c<1$ and $b$ is something small.

This procedure helps us to not blow up to infinity and always improve or stay still. The Armijo rule gives us guarantees that we will always reach a local minimum.

Let's create a function that performs a Newton with both regularization and line search:

<!-- grade_id: cell-62c77a4c9c30838c, points: 0, grade: false, solution: true -->
```python
def armijo_newton_step(x, beta = 1., b = 0.1, c = 0.5):
    ### TO-DO: Implement Damped + Armijo Newton step
    ### Compute Δx (store it in a variabe named `dx`)
    ### ANSWER: Insert code here
    dx = damped_newton_step(x, beta) - x
    a = 1.
    while(f(x + a * dx) > f(x) + b * a * df_dx(x) @ dx):
        a = c * a
    ### END of ANSWER
    return x + a * dx
```

<!-- grade_id: cell-db2a800310cea558, points: 3, grade: true, solution: false -->
```python
np.random.seed(10)

x = np.random.randn(2, 1)
xn = armijo_newton_step(x)
assert(np.isclose(xn, np.array([[-1.07868968, 0.72390241]]).T, rtol=1e-3, atol=1e-3).all())

x = np.random.randn(2, 1)
xn = armijo_newton_step(x)
assert(np.isclose(xn, np.array([[-2.89066819, -0.00831666]]).T, rtol=1e-3, atol=1e-3).all())

x = np.random.randn(2, 1)
xn = armijo_newton_step(x)
assert(np.isclose(xn, np.array([[1.95634526, -0.72452386]]).T, rtol=1e-3, atol=1e-3).all())

np.random.seed(None)
```

```python
# Let's Optimize!
# Plot the function!
plt.close()
fig = plt.figure()
ax = fig.add_subplot(1, 1, 1)

N = 80
x1 = np.linspace(-3., 3., N)
x2 = np.linspace(-3., 3., N)

X, Y = np.meshgrid(x1, x2)

val = np.zeros((N, N))
for i in range(N):
    for j in range(N):
        xx = np.zeros((2, 1))
        xx[0] = X[i, j]
        xx[1] = Y[i, j]
        val[i, j] = f(xx)

CS = ax.contour(x1.reshape((N,)), x2.reshape((N,)), val)

# Initial point
x0 = np.array([[1., 0.7]]).T
# x0 = np.array([[1., 1.25]]).T # It was failing with Damped alone! Now it finds the closest local minima!

ax.plot(x0[0, 0], x0[1, 0], 'rx')

print("Initial value:", f(x0))

x = np.copy(x0)
```

```python
# iterate
x = armijo_newton_step(x)

print("x:", x.T, "->", f(x))

ax.plot(x[0, 0], x[1, 0], 'rx')

fig # show figure again with updated point(s)
```

```python

```